SntB triggers the antioxidant pathways to regulate development and aflatoxin biosynthesis in Aspergillus flavus

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Recommendations For The Authors):

Inclusion of other catalase, peroxidase or superoxide dismutase gene promoters (with ChiP-seq screen shots) and whether they contain sntB binding sites is important to provide other potential downstream pathways controlling oxidative stress mediated regulation of development and aflatoxin metabolism. This can be presented as supplementary material.

or

Some more examples of ChiP-seq peaks in the promoters of nsdC, nsdD, sclR, steA, wetA, veA, fluG, sod2, catA, catC would strengthen the paper for the reliability of the ChiP-seq data. Currently, visualisation of the ChIP-seq data is only limited to catC gene promoter, where background ChIP-seq signals are very high (Figure 5F).

The binding region and motif of SntB on the catA, catB, sod1, and sod2 genes were shown in Figure S7 and described in lane 531-536 and 881-884. The background of ChIP-seq signals is high, but the enrich level in the ip-sntB-HA samples is significant compared to IP-WT.

Figure 5F, letters are too small, and difficult to read. The same is true for Figure 4. Letters should be enlarged for the readers to read it without problem.

Thanks. We have revised the Figure 5F and Figure 4. Please see these Figures.

Reviewer #2 (Recommendations For The Authors):

The authors fully addressed my concerns and made appropriate changes in the manuscript. The quality of the manuscript is now improved.

Thanks. We would like to express our sincere gratitude for your affirmation and thoughtful feedback. Your positive comments have been extremely encouraging and have strengthened my confidence in my work. Your time and effort in reviewing my submission are greatly appreciated.

Changes from the earlier version to the revised version

Abstract

Abstract

The epigenetic reader SntB was identified as an important transcriptional regulator of growth, development, and secondary metabolite synthesis in Aspergillus flavus. However, the underlying molecular mechanism is still unclear. In this study, sntB gene deletion (ΔsntB), complementary (Com-sntB), and HA tag fused to sntB (sntB-HA) strains were constructed by using the homologous recombination method, respectively. Our results revealed that deletion of sntB inhibited the processes of mycelia growth, conidial production, sclerotia formation, aflatoxin synthesis, and ability to colonize host compared to wild type (WT), and the defective phenotype of knockout strain ΔsntB could be restored in its complementary strain Com-sntB. Chromatin immunoprecipitation sequencing (ChIP-seq) of sntB-HA and WT and RNA sequencing (RNA-seq) of ΔsntB and WT strains revealed that SntB played key roles in oxidative stress response of A. flavus. The function of catC (encode a catalase) gene was further analyzed based on the integration results of ChIP-seq and RNA-seq. In ΔsntB strain, the relative expression level of catC was significantly higher than in WT strain, while a secretory lipase encoding gene (G4B84_008359) was down-regulated. Under the stress of oxidant menadione sodium bisulfite (MSB), the deletion of sntB obvious down-regulated the expression level of catC. After deletion of catC gene, the mycelia growth, conidial production, and sclerotia formation were inhibited, while ROS level and aflatoxin production were increased compared to the WT strain. Results also showed that the inhibition rate of MSB to ΔcatC strain was significantly lower than that of WT group and AFB1 yield of the ΔcatC strain was significantly decreased than that of WT strain under the stress of MSB. Our study revealed the potential machinery that SntB regulated fungal morphogenesis, mycotoxin anabolism, and fungal virulence through the axle of from SntB to fungal virulence and mycotoxin bio-synthesis, i.e. H3K36me3 modification-SntB-Peroxisomes-Lipid hydrolysis-fungal virulence and mycotoxin bio-synthesis. The results of the study shed light into the SntB mediated transcript regulation pathway of fungal mycotoxin anabolism and virulence, which provided potential strategy for control the contamination of A. flavus and its aflatoxins.

Introduction

Aspergillus flavus is one of the common asexual species, a saprophytic fungus and the second largest pathogenic fungus after Aspergillus fumigatus, widely distributed in soil, air, water, plants, and agricultural products in nature [1]. Aflatoxins produced by A. flavus have strong toxicity, and is extremely harmful to human society. Animal and human health can be negatively affected by aflatoxins, which are carcinogenic, teratogenic, and mutagenic [2]. Among aflatoxins, AFB1 is the most frequently occurring and the most toxic and carcinogenic, which is converted to AFB1-8 and 9-epoxide in the liver and formed adducts with the guanine base of DNA and thus results in acute and chronic diseases in both human and household animals [3]. According to the Food and Agriculture Organization of the United Nations, 25% of the food crops in the world are contaminated with aflatoxins [4], which are often detected in grains, nuts, and spices [5, 6]. It is urgent to control the contamination of A. flavus and its main mycotoxin, AFB1.

In recent decades, the biosynthetic pathway of aflatoxins was investigated in detail benefit from the sequence of A. flavus genome [7]. This pathway consists of a complex set of enzymatic reactions [8–11]. In general, these enzymes are encoded by clusters of genes, which are regulated by cluster-specific genes: aflR and aflS [12, 13]. The initial stage of aflatoxins biosynthesis is catalyzed by polyketide synthase (PKSA) to form the polyketone backbone [14]. The synthesis of aflatoxins is additionally influenced by environmental stimuli such as pH, light exposure, nutrient availability, and the response to oxidative stress, potentially leading to the alteration of gene expression related to toxin biosynthesis [15–17]. Besides the biosynthetic pathway and its internal gene regulation, protein post-translational modifications (PTMs), an important mean of epigenetics, represent an important role in the regulation of aflatoxins synthesis, including 2-hydroxyisobutyrylation, succinylation, acetylation, and methylation [18–24], in which Snt2 (also called sntB, an epigenetic reader) is deeply involved. Despite advancements in the field, our understanding of the molecular mechanisms of aflatoxin production in A. flavus is still fragmentary.

The epigenetic reader encoded by sntB in A. nidulans was identified as a transcriptional regulator of the sterigmatocystin biosynthetic gene cluster and deletion of sntB gene in A. flavus results in loss of aflatoxin production [25], increasing global levels of H3K9K14 acetylation and impairing several developmental processes [20]. The homolog gene in yeast, SNTB coordinates the transcriptional response to hydrogen peroxide stress [26, 27]. In Penicillium expansum, SntB regulated the development, patulin and citrinin production, and virulence on apples [28]. In A. nidulans, SntB combined with a H3K4 histone demethylase KdmB, a cohesin acetyltransferase (EcoA), and a histone deacetylase (RpdA) to form a chromatin binding complex (KERS) and bound to regulatory genes and coordinated fungal development with mycotoxin synthesis [29]. In A. flavus, the KERS complex also consists of KdmB, RpdA, EcoA, and SntB and plays a key role in the fungal development and secondary metabolites metabolism [30]. sntB also regulated the virulence in Fusarium oxysporum, and respiration in F. oxysporum and Neurospora crassa [31, 32].

However, the specific regulatory mechanism of sntB in A. flavus remains unclear. In this study, we identified the regulatory network of sntB by Chromatin immunoprecipitation sequencing (ChIP-seq) and RNA sequencing (RNA-seq), which shed light on its impact on fungal biology.

Experimental Procedures

A. flavus strains, media, and culture conditions

A. flavus Δku70 ΔpyrG was used as a host strain, for genetic manipulations. All strains used in this study are listed in the Table 1. potato dextrose agar (PDA, 39 g/L, BD, Difco, Franklin, NJ, USA), complete medium (CM, 6 g/L tryptone, 6 g/L yeast extract, 10 g/L glucose), and potato dextrose (PDB, 24 g/L, BD, Difco, Franklin, NJ, USA) were used for mycelial growth and sporulation determination, sclerotia production, and mycotoxin production analysis, respectively. All experiments were technically repeated three times and biologically repeated three times.

The construction of mutant strains

All mutant strains, including sntB and catC gene knock-out strain (ΔsntB and ΔcatC), the complementation strain for the ΔsntB strain (Com-sntB), and HA tag fused to sntB strain (sntB-HA), were constructed following the protocol of homologous recombination [34] and the detail protocol was as described in our previous study [35]. The related primers were listed in Table 2. The constructed strains were confirmed by diagnostic PCRs and by southern blot [36]. The construction of sntB-HA was further determined by western blot with anti-HA antibody as descripted previously [35].

Phenotypic analysis and aflatoxin analysis

The spores (107 conidia/mL) of wild type (WT), ΔsntB, and Com-sntB strains were used. The details of experiment were according to our previous study [37]. Hyphal septum was stained according to the descripted method [38]. Each fungal strain was evaluated on four plates, and each experiment was repeated three times.

Fungal colonization on crop kernels

According to our previous experimental protocol [37], the colonizing ability of WT, ΔsntB, and Com-sntB fungal strains on peanut and corn kernels was analyzed. The crop kernels were disinfected with 0.05% sodium hypochlorite and soaked for 30 min in a solution containing 105 conidia/mL fungal spores. Afterwards, the seeds were placed in a petri dish and cultured at 29°C for 6 d. Finally, the number of conidia was calculated and AFB1 product was analyzed by TLC.

Animal invasion assay

The animal invasion assay using silkworms was conducted according to a previous study [34, 39]. Silkworms (Bombyx mori) were randomly separated into four groups (10 larvae/group) when silkworm larva reach about 1 g in weight. Each silkworm was injected with 5 µL saline, or 5 µL conidial suspension (106 spores/mL) from WT, ΔsntB, and Com-sntB strains. The survival rate of silkworms was calculated. Dead silkworms were transferred into fresh 9 cm Petri dishes and inoculated for 5 d in the dark. The conidia number and AFB1 production from each group were measured.

RNA-seq analysis

To reveal the potential complex regulatory network of the sntB, RNA-seq analysis was carried out on the WT and ΔsntB strains. RNA-seq analysis was carried out by APPLIED PROTEIN TECHNOLOGY, Shanghai (www.aptbiotech.com) [40]. Data processing was according to a previous study [41]. Differentially expressed gene (DEGs) were assigned as genes with |log2FoldChange|>1 and adjusted P-adj< 0.05. Gene Ontology (GO) and Kyoto Encyclopedia of Genes and Genomes (KEGG) pathways were used to analyze the functions of DEGs.

ChIP-seq and data analysis

ChIP-seq analysis was carried out on the WT and sntB-HA strains. The conidia (104/ml) for each strain were inoculated in 100 mL PDB shaking at 180 rpm under 29°C for 72h, and subjected to ChIP-seq analysis by Wuhan IGENEBOOK Biotechnology Co.,Ltd (http://www.igenebook.com). ChIP experiment was according to a previous study [34]. Raw sequencing with low-quality reads discarded and reads contaminated with adaptor sequences trimmed were filtered by Trimmomatic (v0.36) [42]. The clean reads were mapped to the reference genome of A. flavus by Burrows-Wheeler Alignment tool (BWA, version 0.7.15) [43]. MACS2 (v2.1.1) and Bedtools (v2.25.0) were used for peak calling and peak annotation, respectively. Differential binding peaks were identified by Fisher’s test with q value <0.05. HOMER (version3) was used to predict motif occurrence within peaks with default settings for a maximum motif length of 12 base pairs [44]. Genes less than 2000 bp away were associated to the corresponding peak. GO and KEGG enrichment analyses of annotated genes were implemented in EasyGO [45] and KOBAS (v2.1.1) [46], with a corrected P-value cutoff of 0.05.

Quantitative RT-PCR (qRT-PCR) analysis

The fungal spores (106/mL) were cultured in PDB medium for 48 h, and then mycelium was ground into powder with liquid nitrogen. Total RNA was prepared by TRIpure total RNA Extraction Reagent (Bestek, China) according to the protocol used by Zhang [47]. qRT-PCR was performed according to previously described [48], and the primers were shown in Table 3.

Oxidative stress assays

To evaluate the role of sntB in fungal resistance to oxidative stress, a series concentration (0, 0.12, 0.24, and 0.36LmM) of menadione sodium bisulfite (MSB) were added to the medium. 106 fungal spores for each strain were inoculated on the medium and cultured in dark at 37L. The diameters of colonies were measured 3 d after inoculation and the inhibition rate was calculated as previously descripted [35]. The AFB1 product was analyzed by TLC after the strains were cultured in YES medium in dark at 29L for 7 d.

Reactive oxygen species assay

With the instructions provided in the user’s manual, the intracellular ROS production was measured using a ROS assay kit (S0033S, Beyotime Institute of Biotechnology, China). After harvest, the mycelia were incubated with 10 µM 6-carboxy-2’,7’-dichlorodihydrofluorescein diacetate (DCFH-DA) and 50 g/mL Rosup for 30 min. With SpectraMax Imaging Cytometer (Molecular Devices, Sunnyvale, Calif.) at emission wavelength of 525 nm and excitation wavelength of 488 nm, fluorescence signals of intracellular ROS production were acquired.

Statistical analysis

All data in this study were expressed as mean ± standard deviation. The statistical analysis was performed using the software GraphPad Prism8 (GraphPad Software, La Jolla, CA, USA). The difference was considered to be statistically significant when P<0.05.

Results

The phenotype of SntB in A. flavus

The role of SntB in A. flavus has been previously characterized by analyzing both ΔsntB and overexpression sntB genetic mutants [20]. To further investigate the intrinsic mechanism of this regulator on the development and aflatoxin biosynthesis in A. flavus, the sntB deletion strain (ΔsntB) and the complementary strain (Com-sntB) were constructed by the method of homologous recombination and verified by diagnostic PCR and southern blotting (Figure S1). The expression levels of sntB in WT, ΔsntB, and Com-sntB strains was further detected by qRT-PCR and the result showed that the expression of sntB was absent in the gene-deletion strains, and it fully recovered in the Com-sntB (Figure S1), which reflected that the sntB and Com-sntB had been constructed, and could be used in the subsequent experiments of this study. The phenotype analysis of this study revealed that the deletion of sntB gene significantly inhibited the growth of mycelium, hyphae morphology, the length of fungal cell (between two adjacent septa), the number of conidiation, sclerotium formation, and aflatoxin biosynthesis, while the above phenotypes of both development and mycotoxin bio-synthesis were fully recovered in the Com-sntB strain (Figure 1). To reveal the signaling pathways of sntB in conidiation, sclerotium formation, and aflatoxin biosynthesis, qRT-PCR analysis was performed to assess the expression levels of sporulation related transcriptional factor genes, steA, wetA, fluG, and veA, sclerotia formation related transcriptional factor genes, nsdC, nsdD, and sclR [49, 50], and the AFs synthesis gene cluster structural genes aflC, aflR, and aflP, and the main regulatory gene aflR and aflS. As shown in Figure S2, the relative expression levels of these genes were significantly lower in the ΔsntB strain compared to that of the WT strain, and recovered in the Com-sntB strain. These results indicated that SntB regulates the conidiation, sclerotium formation, and aflatoxin biosynthesis by the canonical signaling pathways mediated by these regulators.

The functions of SntB in A. flavus. (A) The colonies of WT, ΔsntB, and Com-sntB strains grown on PDA at 37°C in dark for 4 d. (B) The colony diameter statistics of the above fungal strains. (C) Microscopic examination revealed the difference in mycelia of each fungi strain at 37L in dark, scale=200 μm. (D) Microscopic examination of the hyphal septum of each strain at 37L in dark, scale=50 μm. (E) The spore production statistics. (F) All the above fungal strains were point-inoculated on CM medium and grown for 7 d at 37L. (G) The number of sclerotia of the above fungal strains. ND=Not detectable. (H) AFB1 production of the above fungal strains was detected by TLC after the strains incubating at 29L in PDB medium for 7 d.

SntB plays important role in virulence of A. flavus to both plant and animal hosts

In order to explore the effect of SntB on the fungal colonization ability, peanut and maize kernels were infected with spore solution of each fungal strain. Compared with WT, the conidiation yield of ΔsntB on the infected host was significantly reduced (P<0.001) and no AFB1 could be detected on the hosts infected by ΔsntB, while in the Com-sntB strain, the capacity to produce conidia and AFB1 on both crop kernels recovered (Figure 2A-2C and Figure S3A). The role of sntB in fungal virulence to animals was also investigated. As Figure 2D-2E shown, the survival rate of silkworms injected by spores of ΔsntB strain was significantly higher than that from WT infected larvae. There was less fungal mycelium, conidia, and AFB1 production on the dead silkworms injected by ΔsntB compared to the silkworms from the WT injection groups, but when the gene was reintroduced (i.e. the Com-sntB group), similar to what found in the WT group, the survival rate of silkworms obviously dropped and more fungal mycelium, conidia, and AFB1 produced on the dead silkworms (Figure 2F-2H). All the above results revealed that SntB plays an essential role in virulence of A. flavus.

The role of SntB on the ability of A. flavus to colonize host. (A) Phenotypic of corn and peanut kernels colonized by ΔsntB, Com-sntB, and WT strains at 29°C in dark for 7 d. (B) Statistical of the number of conidia on the surface of peanut and maize kernels. (C) TLC analysis to detect the yield of AFB1 in kernels infected by the above fungal strains after 7 d incubation. (D) Photographs of the silkworms infected by the above fungal strains. (E) The survival rate of silkworms 1 week after injection of the above strains. (F) Photographs of the dead silkworms infected by A. flavus after 6 d incubation. (G) The spore production statistics of the above fungal strains on the dead silkworms shown in (F). (H) TLC analysis of AFB1 levels produced in infected dead silkworms in (F).

The capacity of fungal infection is closely related to secreted hydrolases, such as amylase, lipase, protease, etc. In order to explore the effect of SntB on the activity of hydrolases, the activities of amylase in the above each fungal strain were further determined. The results showed that the colonies of ΔsntB produced almost no degradation transparent circle after adding iodine solution compared with that of WT and Com-sntB, which indicated that the activity of α-amylase in ΔsntB strain was significantly reduced (P<0.001) (Figure S3B and S3C). This suggests that SntB plays an important role in the fungal pathogenicity by changing the hydrolases activity of A. flavus.

SntB chords global gene expression

To explore the downstream signaling pathways regulated by SntB, samples with three biological replicates of WT and ΔsntB strains were submitted for RNA-Seq. In the assay, the bases score Q30 was more than 93.19% (Table S1) and mapping ratio was from 95.17% to 95.80% (Table S2). To further confirm the quality of RNA-Seq, PCA (Principal Components Analysis) and Pearson correlation analysis were performed. Correlation analysis revealed that the samples were clustered by groups (Figure 3A). A plot of PC1 (47.70%) and PC2 (23.20%) scores showed a clear separation between the groups (Figure 3B). A total of 1,446 and 1,034 genes were significantly up- and down-regulated, respectively, in the ΔsntB compared to the WT strain (Figure 3C and Supplementary Table S3). GO enrichment analysis identified 93 enriched GO terms (P<0.05) (Figure 3D and Supplementary Tables S4). In the biological process category, the most enriched terms were “oxidation-reduction process (GO: 0055114)”. In the molecular function category, “catalytic activity (GO: 0003824),” “oxidoreductase activity (GO: 0016491),” and “cofactor binding (GO: 0048037)” were the most significantly enriched terms. Whereas terms associated with “Set3 complex (GO: 0034967)”, “mitochondrial crista junction (GO: 0044284)” and “extracellular region (GO: 0005576)” were significantly enriched in the cellular component category. Additionally, all the DEGs were mapped according to the KEGG database, and 42 significantly enriched pathways were identified (P<0.05) (Supplementary Table S5). Among them, “metabolic pathways (ko01100),” “aflatoxin biosynthesis (ko00254),” and “microbial metabolism in diverse environments (ko01120)” were the most significantly enriched (Figure 3E and Supplementary Table S5).

SntB chords global gene expression in A. flavus. (A) The Pearson correlation results shown by heatmap. (B) Principal component analysis (PCA) on six fungal samples, including three ΔsntB (KOsnt2) and three WT fungal strains. (C) Volcano map reflecting the distribution of the differential expression genes. (D) Gene Ontology (GO) analyses of the differentially expressed genes. (E) Kyoto encyclopedia of genes and genomes (KEGG) analyses of the differentially expressed peaks related genes.

Characterization the binding regions of SntB by chromatin immunoprecipitation sequencing (ChIP-seq)

To characterize the chromatin regions targeted by SntB, ChIP-seq studies were carried out with both HA tag fused sntB strain (sntB-HA) and WT strain. The sntB-HA strain was constructed by homologous recombination through fused HA to the 3’ end of sntB (Figure 4A). In the ChIP-seq assay, more than 94.66% of bases score Q30 and above in each sample (Table S6), and reaching 52.50% to 94.48% of mapping ratio (Table S7). The principal component analysis (PCA) (Figure S4A) and heat map (Figure S4B) reflected that the quality of samples was competent for subsequent ChIP-seq analysis. There were 1,510 up-enriched differently accumulated peaks (DAPs) in sntB-HA fungal strain compared to the WT strain, which were distributed on the whole A. flavus genome (Figure 4B, and Table S8). Most of the up-enriched peaks were located in the promoter (82.85%) region (Figure 4C). To determine binding regions of SntB, we used the HOMER known and de novo motif discovery algorithm. Motifs were sorted based on p-values and the top 5 enriched known motifs were shown in Figure 4D. The results consisted of motifs derived from previously-published ChIP-Seq experiments on Cbf1, bHLHE40, NFY, Usf2, and USF1 (Table S9). However, the most enriched de novo motif was NFYA (1e-97) (Figure 4E). The genes of the DAPs were further subjected to GO and KEGG analysis. The most strikingly enriched GO terms in the biological process category were “cell communication (GO:0007154)”, “response to stimulus (GO:0050896)”, and “response to external stimulus (GO:0009605)”. Whereas terms associated with “DNA-binding transcription factor activity (GO:0003700)”, “DNA-binding transcription factor activity, RNA polymerase II-specific (GO:0000981)”, and “sequence-specific DNA binding (GO:0043565)” were the most significantly enriched molecular function category (Figure 4F and Table S10). And these genes were mostly enriched in “Methane metabolism” and “MAPK signaling pathway - yeast” pathways (P value <0.05) (Figure 4G and Table S11).

Characterization the binding regions of SntB. (A) Verification of the construction of sntB-HA strain using western blot. M means the PAGE-MASTER Protein Standard Plus (GenScript USA, MM1397). (B) The distribution of differently accumulated peaks on the genome. (C) Vennpie map of the differently accumulated peaks distribution on gene functional elements. (D) Enrichment of known motifs showing the top-ranked motif logos. (E) Enrichment of de novo motifs showing the top-ranked motif logos. (F) GO analyses of the of the differently accumulated peaks related genes. (G) KEGG analyses of the differently accumulated peak related genes.

Integration of the results of ChIP-seq and RNA-seq assays

Earlier version: After the overlapping the results from both different sequence methods (ChIP-seq and RNA-seq), 238 DEGs were found (Figure 5A). According to the GO annotation, these DEGs were significantly enriched in 8 GO terms, including “cellular response to reactive oxygen species (GO:0034614)”, “reactive oxygen species metabolic process (GO:0072593)”, and “cellular response to oxygen-containing compound (GO:1901701)” (Figure 5B). It was further noted that the DEGs were significantly assigned to “Carbon metabolism (afv01200)”, “Peroxisome (afv04146)”, and “Glyoxylate and dicarboxylate metabolism (afv00630)” KEGG pathways (Figure 5C and Table S12-S13). These results revealed that SNTB is essential for A. flavus to maintain the homeostasis of intracellular reactive oxygen species. Studies had shown that SNTB could response to oxidative stress in yeast [27] and Magnaporthe oryzae [51]. As Figure 5D and Figure 5E showed, due to the deletion of the sntB gene, ΔsntB exhibited a severe MSB sensitivity phenotype compared to that of the wild-type (WT) strain, and the phenotype recovered in the complementary strain (Com-ΔsntB). The results showed that the inhibition rate of oxidant MSB to ΔsntB would be significantly enhanced with the increase of MSB concentration. These results showed that SntB deeply participates in the regulation of oxidative stress pathway. As the most abundant peroxisomal enzyme, catalases catalyze decomposition of hydrogen peroxide [52]. As the most abundant peroxisomal enzyme, catalases catalyze decomposition of hydrogen peroxide [52]. To further study the mechanism of SntB mediated oxidative response of A. flavus, the catC (encode a catalase) gene was selected based on the above integration results. According to the peak map, the binding region of SntB on catC gene was significantly enriched in sntB-HA strains compared to WT strain (Figure 4F). The motif in the binding region was shown in Figure 4G and the sequence was TCCGCCCG. The relative expression levels of catC in wild-type and ΔsntB strains under MSB treatment were measured. As Figure 5H shown, the expression level of catC was significantly higher in ΔsntB strain than in WT strain, it suggested that to compensate the absence of sntB, catC is up-regulated to respond the higher intracellular oxidative level. However, under the stress of oxidant MSB, the deletion of sntB obvious suppressed the expression level of catC compared to that of WT strain, which reflected that the absence of sntB significantly impaired the capacity of catC to further respond to extra external oxidative stress. These results revealed that SntB is deeply involved in CatC mediated oxidative response in A. flavus.

Revised version: After the overlapping the results from both different sequence methods (ChIP-seq and RNA-seq), 238 DEGs were found (Figure 5A). According to the GO annotation, these DEGs were significantly enriched in 8 GO terms, including “cellular response to reactive oxygen species (GO:0034614)”, “reactive oxygen species metabolic process (GO:0072593)”, and “cellular response to oxygen-containing compound (GO:1901701)” (Figure 5B). It was further noted that the DEGs were significantly assigned to “Carbon metabolism (afv01200)”, “Peroxisome (afv04146)”, and “Glyoxylate and dicarboxylate metabolism (afv00630)” KEGG pathways (Figure 5C and Table S12-S13). These results revealed that SNTB is essential for A. flavus to maintain the homeostasis of intracellular reactive oxygen species. Studies had shown that SNTB could response to oxidative stress in yeast [27] and Magnaporthe oryzae [51]. As Figure 5D and Figure 5E showed, due to the deletion of the sntB gene, ΔsntB exhibited a severe MSB sensitivity phenotype compared to that of the wild-type (WT) strain, and the phenotype recovered in the complementary strain (Com-ΔsntB). The results showed that the inhibition rate of oxidant MSB to ΔsntB would be significantly enhanced with the increase of MSB concentration. These results showed that SntB deeply participates in the regulation of oxidative stress pathway. As the most abundant peroxisomal enzyme, catalases catalyze decomposition of hydrogen peroxide [52]. As the most abundant peroxisomal enzyme, catalases catalyze decomposition of hydrogen peroxide [52]. To further study the mechanism of SntB mediated oxidative response of A. flavus, the catC (encode a catalase) gene was selected based on the above integration results. According to the peak map, the binding region of SntB on catC gene was significantly enriched in sntB-HA strains compared to WT strain (Figure 5F). The motif in the binding region was shown in Figure 5G and the sequence was TCCGCCCG. The relative expression levels of catC in wild-type and ΔsntB strains under MSB treatment were measured. As Figure 5H shown, the expression level of catC was significantly higher in ΔsntB strain than in WT strain, it suggested that to compensate the absence of sntB, catC is up-regulated to respond the higher intracellular oxidative level. However, under the stress of oxidant MSB, the deletion of sntB obvious suppressed the expression level of catC compared to that of WT strain, which reflected that the absence of sntB significantly impaired the capacity of catC to further respond to extra external oxidative stress. These results revealed that SntB is deeply involved in CatC mediated oxidative response in A. flavus.

Integration of the results of ChIP-seq and RNA-seq assays. (A) Venn diagrams of ChIP-seq and RNA-seq. (B) GO analyses of the common genes. (C) KEGG analyses of the common genes. (D) The phenotype of WT, ΔsntB, and Com-sntB strains cultured in PDA containing a series concentration of MSB for 3 d. (E) Statistical analysis of the growth inhibition rate of MSB to all the above fungal strains according to Panel D. (F) Comparation of the enrich levels of the SntB binding region of catC gene between WT and sntB-HA strains. (G) The motif logo in the SntB binding region of catC gene. (H) The relative expression level of sod1 in WT and ΔsntB strains with or without MSB treatment.

CatC is important for A. flavus response to oxidative stress

The functions of catC gene in A. flavus were further explored by knockout of the catC (Figure S5). As shown in Figure 6A-6C, the diameter of ΔcatC strain was significantly smaller than that of WT, and the conidia number in the ΔcatC strain decreased significantly compared to that of WT. The sclerotia production of ΔcatC strain was also significantly less than that of WT strain (Figure 6D and 6E). In view of CatC is involved in the oxidative response pathways of A. flavus (Figure 4), both ΔcatC and WT strains were treated by a serial concentration of MSB, and the results showed that the inhibition rates of MSB to ΔcatC strain was significantly lower than that of WT group (Figure 6F and 6G). Catalase is a major peroxisome protein and plays a critical role in removing peroxisome-generated reactive oxygen species. The result of fluorescence intensity of oxidant-sensitive probe dichlorodihydrofluorescein-diacetate (DCFH-DA) showed that ROS accumulation in the ΔcatC strain was higher than that in the WT strain (Figure 6G). This result echoed that the deletion of sntB increased intracellular oxidative level and the inhibition rate of MSB, and up-regulated the expression of catC (Figure 5D-5F). The role of CatC in the bio-synthesis of AFB1 was also assessed (Figure 6H). The results showed that a relatively large amount of AFB1 was produced by the ΔcatC strain compared to the WT. But when under the stress of MSB, AFB1 yield of the WT strain was significantly more than that of ΔcatC strain. All the above results revealed that the catC plays an important role in SntB mediated regulation pathway on fungal morphogenesis, oxidative stress responding and AFB1 production.

The functions of catC in A. flavus. (A) The colonies of WT and ΔcatC strains grown on PDA at 37°C in dark for 4 d. (B) The colony diameter statistics of the above fungal strains. (C) The spore production statistics of the above fungal strains. (D) All above fungal strains were point-inoculated on CM medium and grown for 7 d at 37L. (E) The number of sclerotia of the above fungal strains. (F) The phenotype of above strains cultured in PDA medium containing a series concentration of MSB for 3 d. (G) Statistical analysis of the growth inhibition rate of MSB to all the above fungal strains according to (F). (H) Relative ROS level in the WT, ΔsntB, Com-sntB, and ΔcatC strains. (I) AFB1 production of the above fungal strains was detected by TLC after the strains incubating at 29L in PDB medium for 7 d.

SntB regulates fungal virulence through peroxidase mediated lipolysis

Biogenesis of peroxisomes was reported to promote lipid hydrolysis, produce the production of glycerol, and further change fungal pathogenicity [53]. Due to deeply involved in oxidative response of A. flavus, we wondered if SntB also takes part in the regulation of the production of lipid and glycerol. As shown in Table S14, one gene (G4B84_008359) in lipase activity GO term was significantly down-regulated in ΔsntB strain, which encoded a secretory lipase belonged to the virulence factors reported in Pseudomonas aeruginosa [54]. The lipase activity was assayed by examining the ability to cleave glycerol tributyrate substrate [55]. The results showed that the colony diameter of ΔsntB strain on PDA medium with tributyrin were significantly smaller than that of the control, and the colony diameters of WT and Com-sntB strains on PDA medium were obviously bigger than those on 0.3% tributyrin PDA medium. The relative inhibition rate of tributyrin to the colony growth of ΔsntB strains was significantly higher than that of WT and Com-sntB strain (Figure S3D and S3E). Our previous study revealed that H3 lysine 36 trimethylation (H3K36me3) modification on the chromatin region of the sntB is regulated by AshA and SetB [34]. H3K36me3 usually promote gene transcription [56, 57]. Our study revealed the potential machinery associated with SntB mediated regulation on fungal morphogenesis, mycotoxin anabolism and fungal virulence, which lurks the axle of from SntB to fungal virulence and mycotoxin bio-synthesis through lipid catabolism (i.e. H3K36me3 modification-SntB-Peroxisomes-Lipid hydrolysis-fungal virulence and mycotoxin bio-synthesis).

Discussion

SntB is a conserved regulator in many species, including A. nidulans [25], Saccharomyces cerevisiae [27], Schizosaccharomyces pombe [58], F. oxysporum [31], Neurospora crassa [32], and A. flavus [20, 25]. SntB can regulate the production of uncharacterized secondary metabolites, including aspergillicin A1 and aspergillicin F2 [59]. MoSntB protein was required for regulation of infection-associated autophagy in M. oryzae [51]. In A. flavus, the functions of sntB gene were analyzed by both ΔsntB and overexpression sntB genetic mutants[20]. SntB deletion impaired several developmental processes, such as sclerotia formation and heterokaryon compatibility, secondary metabolite synthesis, and ability to colonize host seeds, which were consistent with our results (Figure 1 and 2). Unlike, a complementation strain was constructed in this study which further clarified and confirmed the function of sntB gene. In this study, the potential mechanism under these effects were further analyzed by detection the related transcriptional factor genes of sporulation (steA, WetA, fluG, and veA), sclerotia formation (nsdC, nsdD, and sclR), and the AFs synthesis related genes aflC, aflD, aflO, aflP, and aflR (Figure S2). In the RNA-seq data, we also found some DEGs related to AFs synthesis (aflB, aflE, aflH, aflK, aflN, aflO, aflP, aflQ, aflR, aflS, aflV, and aflW) (Figure S6). And all these genes were down-regulated, which was consistent with that the AFs production in ΔsntB was significantly decreased compared to WT and Com-sntB (Figure 1F). These results inferred that SntB regulated the morphogenesis and the production of A. flavus through the above canonical signal pathways.

For the process of A. flavus invading hosts, in view of it is a notorious pathogen for plant and animal, we established both crop and insect models, especially silkworm represented animal mode profoundly revealed the critical role of SntB in fungal virulence (Figure 2). The results of crop kernel models showed that the number of spores of ΔsntB on kernels of both peanut and maize was dramatically lower than that of strains WT and Com-sntB (Figure 2A and 2B) and almost no AFB1 was detected on maize and peanut kernels inoculated with ΔsntB, while plenty of AFB1 were detected from the kernels infected by WT fungal strain, and AFB1 biosynthesis capacity of Com-sntB strains recovered compared to the ΔsntB and WT fungal strains (Figure 2C). These results were corroborated by previous study [20]. It was also found in this study that the survival rate of silkworms injected by spores of ΔsntB strain was significantly higher than the silkworms injected with spores from WT and Com-sntB fungal strains (Figure 2D and 2E). What’s more, we also assayed the effect of SntB on the activity amylase, which was closely related to the capacity of fungal infection [60]. As shown in Figure S3C, after adding iodine solution, the ΔsntB strain almost did not produce a degradation transparent ring compared with wild-type and complementary strains, indicating the amylase activity of the sntB gene knockout strain was significantly decreased (P<0.001). Our results comprehensively reveal the important function of SntB in the growth, development, secondary metabolite synthesis, and virulence of A. flavus.

SntB was reported as an important epigenetic reader. In A. flavus, SntB was reported to regulate global histone modifications (acetylation and methylation) and interact with EcoA and RpdA to form a conserved chromatin regulatory complex [20, 29]. Loss of sntB in M. orzyae also led to an increase in H3 acetylation [51]. In our RNA-seq data, we also found a set domain containing histone-lysine N-methyltransferase (Ash1, G4B84_009862) was down-regulated in ΔsntB strain compared to WT (Table S3), which was reported to regulate mycotoxin metabolism and virulence via H3K36 methylation in A. flavus [34]. Besides, SntB is reported to be a transcriptional regulator in A. nidulans [25] and F. oxysporum [31]. So, we use RNA-seq and ChIP-seq to study the transcriptional response of sntB in A. flavus. By integration analysis the results of ChIP-seq and RNA-seq assays, we found that the enriched GO terms and KEGG pathways of the DEGs were related to oxidative response (Figure 5A-5C). These results reflected that SntB plays an important role in fungal responding to oxidative stress, which is consistent with previous reports that SntB could response to oxidative stress in yeast [27], F. oxysporum [31], and M. oryzae [51].

As a harmful by-product of oxidative metabolism, reactive oxygen species (ROS) is unavoidable and essential for fungus development [61, 62]. ROS has also been shown to be required for aflatoxin production [63–65]. Several oxidative stress-responsive transcription factors have been identified as regulating aflatoxin production, including AtfB, AP-1, and VeA [64, 66–68]. Previous studies shown that SntB protein coordinates the transcriptional response to hydrogen peroxide-mediated oxidative stress in the yeast [26, 27] and is involved in fungal respiration and reactive oxidative stress in F. oxysporum and Neurospora crassa [31, 32]. Several GO terms (“cellular response to reactive oxygen species”, “reactive oxygen species metabolic process”, and “cellular response to oxygen-containing compound”) and KEGG pathways (Peroxisome) were enriched by the DEGs screened out by integration of ChIP-seq and RNA-seq data in this study (Figure 5). And the intracellular ROS level in the ΔsntB and ΔcatC strains were significantly higher than that in WT strain (Figure 6H), which was similar to previous report on the cat1 gene in A. flavus [69]. This is the first time to show that the SntB in A. flavus is important in oxidative stress response, through which SntB participates the regulating of aflatoxin bio-synthesis and fungal development.

Earlier version: Fungal defense against ROS is mediated by superoxide dismutases (SOD), catalases (CAT), and glutathione peroxidases (GPX). The effect of MSB on cellular growth and antioxidant enzyme induction in A. flavus was previously explored [70–72]. Once in the cell, menadione may release superoxide anion [73], which was scavenged by superoxide dismutases (SOD) and transformed into hydrogen peroxide, or react with nitric oxide to form peroxynitrite [74]. This study found that after knock out sntB gene, the strain growth was significantly inhibited by MSB (Figure 5F and 5G). Some genes encoded superoxide dismutases and catalases were reported to associated with AF/ST synthesis [75], including mnSOD, sod1, sod2, catA, catB, and hyr1. In our RNA-seq data, 7 related genes were screened out (Table S15), among which, bZIP transcription factor Atf21 (G4B84_008675), catlase C (G4B84_009242), catlase A (G4B84_010740), superoxide dismutases sod2 (G4B84_003204), peroxisomal membrane protein PmpP24 (G4B84_001452) were up-regulated, while catlase B (G4B84_008381) and superoxide dismutase sod1 (G4B84_009129) were down-regulated in ΔsntB strain, respectively.

Revised version: Fungal defense against ROS is mediated by superoxide dismutases (SOD), catalases (CAT), and glutathione peroxidases (GPX). The effect of MSB on cellular growth and antioxidant enzyme induction in A. flavus was previously explored [70–72]. Once in the cell, menadione may release superoxide anion [73], which was scavenged by superoxide dismutases (SOD) and transformed into hydrogen peroxide, or react with nitric oxide to form peroxynitrite [74]. This study found that after knock out sntB gene, the strain growth was significantly inhibited by MSB (Figure 5F and 5G). Some genes encoded superoxide dismutases and catalases were reported to associated with AF/ST synthesis [75], including mnSOD, sod1, sod2, catA, catB, and hyr1. In our RNA-seq data, 7 related genes were screened out (Table S15), among which, bZIP transcription factor Atf21 (G4B84_008675), catlase C (G4B84_009242), catlase A (G4B84_010740), superoxide dismutases sod2 (G4B84_003204), peroxisomal membrane protein PmpP24 (G4B84_001452) were up-regulated, while catlase B (G4B84_008381) and superoxide dismutase sod1 (G4B84_009129) were down-regulated in ΔsntB strain, respectively. The binding region of SntB on catA, catB, sod1, sod2, and catC genes promoter were significantly enriched in sntB-HA strains compared to WT strain (Figure 5F and Figure S7A). Also, the motif in the binding regions were shown in Figure 5G and Figure S7B. These findings suggest that SNTB has the ability to interact with the promoter region of genes associated with oxidative processes, thereby modulating the expression of these genes.

Some studies reported the correlation between ROS formation, aflatoxin production, and antioxidant enzyme activation. Aflatoxin B1 biosynthesis and the activity of total superoxide dismutase were effectively inhibited by cinnamaldehyde, whereas the activities of catalase and glutathione peroxidase were opposite [76]. The expression of catA, cat2 and sod1, and CAT enzymatic activity were opposite correlated to AFB1 biosynthesis under AFs inhibitor piperine treatment [77]. Deletion the gene of sod (GenBank accession no: CA747446) reduced AFs production [78], which was most similar to sod2 (G4B84_003204). The mitochondria-specific sod and the genes aflA, aflM, and aflP belonging to the AFs gene cluster were reported to co-regulated [79]. Ethanol can inhibit fungal growth and AFB1 production in A. flavus and enhanced levels of anti-oxidant enzymatic genes including Cat, Cat1, Cat2, CatA, and Cu, Zn superoxide dismutase gene sod1. All these reports indicated that the expression of anti-oxidant enzymatic genes was opposite correlated to AFB1 biosynthesis.

In our study, 7 genes related to oxidative response were obviously differentially expressed in transcriptome data (Table S15). Among these DEGs, 5 out of 7 genes were up-regulated in ΔsntB strain. Based on the AFs production in ΔsntB was significantly decreased compared to WT and Com-sntB (Figure 1F), the most up-regulated gene in ΔsntB strain, catC (G4B84_000242), was selected for further analysis. We found that the deletion of sntB significantly up-regulated the catC gene, however, the expression of the catC gene was almost undetectable under MSB treatment (Figure 5F). Results also showed that the inhibition rates of MSB to ΔcatC strain was significantly lower than that of WT group and AFB1 yield of the ΔcatC strain was significantly decreased than that of WT strain under the stress of MSB (Figure 6F-6H). These results indicated that SntB is profoundly involved in the catC mediated oxidative response.

Peroxisomes are intimately associated with the metabolism of lipid droplets [80] and the histone lysine methyltransferase ASH1 promotes peroxisome biogenesis, inhibits lipolysis, and further affects pathogenesis of Metarhizium robertsii [53]. Set2 histone methyltransferase family in A. flavus, AshA and SetB, were found to regulate mycotoxin metabolism and virulence via H3K36me3, including the chromatin region of the sntB [34]. By ChIP-seq and RNA-seq, SntB was found to essential for A. flavus to maintain the homeostasis of intracellular reactive oxygen species (Figure 5A) and several anti-oxidant enzymes were up-regulated in ΔsntB strain (Table S14). In addition, we also found only one down-regulated DEG (G4B84_008359) in lipase activity GO term in our RNA-seq data (Table S14), which encodes a secretory lipase and belongs to the virulence factors reported in Pseudomonas aeruginosa [54]. These results suggested that SntB plays a pivotal role in regulating peroxisome biogenesis to promote lipolysis involving in fungal pathogenesis.

Overall, we explored and clarified the bio-function of the SntB and found that SntB responses to oxidative stress through related oxidoreductase represented by CatC in A. flavus (Figure 7). Our study revealed the potential machinery associated with SntB mediated regulation on fungal morphogenesis, mycotoxin anabolism and fungal virulence, which lurks the axle of from SntB to fungal virulence and mycotoxin bio-synthesis (i.e. SntB -Peroxisomes-Lipid hydrolysis-fungal virulence and mycotoxin bio-synthesis). The work of this study provided a novel perspective for developing new prevention and control strategies against pathogenic fungi.

SntB regulate peroxisome biogenesis, fatty acid utilization, and fungal pathogenicity in A. flavus. (A) The phenotype of each strain on PDA medium containing 0.3% tributyrin, (B) Statistics of inhibition rates. The asterisk *** above the bars represents significantly different (p<0.001). (C) Mechanistic diagram of the bio-functions of SntB in A. flavus.

Removed in the revised version: Supporting information